# EEN060 - Applied object-oriented programming

Teacher: [Carlos Natalino](https://www.chalmers.se/en/staff/Pages/Carlos-Natalino-Da-Silva.aspx) / Examiner: [Paolo Monti](https://www.chalmers.se/en/staff/Pages/Paolo-Monti.aspx)

[Course channel on Chalmers Play](https://play.chalmers.se/channel/EEN060_EEN065_Applied_object-oriented_programming/300149)

Before you turn this assignment list in, make sure everything runs as expected.
First, **restart the kernel** and then **run all cells**.
Then, check if all the tests run correctly.
Note that if one of the problems present an error, the following ones **will not** be tested.

In case of discrepancies between the problem command and the tests, you should solve it having in mind the tests.

There are two types of cell:
1. *solution cells:* These are the cells where you write your answer, or modify the existing code to solve the problem.
2. *test cells:* These cells are used to test whether your solution is correct or not. If the tests run correctly, you should see a message `tests passed`. Otherwise, you should see an error message.

**Delete** the line `raise NotImplementedError()` from the problems that you solve.

**Do not delete or add any cell in this file.** All cells that you need are already in place.

You may also have problems that will be manually graded (e.g., HTML/CSS/JS/pseudocode problems).
These problems **do not** have a respective *test cell*.

If you want to execute a cell, select the cell and press **CTRL+Enter** (in Windows) or **CMD+Enter** (in macOS) or click on the **Run cell** button.

**If you want, you can solve this programming assignment using Google Colab**

Link: https://colab.research.google.com/

Just uncomment the following line to do so.

---

# Assignment 0

In this assignment, students should practice:
- how to download the programming assignments
- how to solve the programming assignments
- how to upload them back to canvas
- how to solve the Python problems
- how to debug your code
- how to solve the HTML problems

**Observations:**
For the problems in this assignment, you should:
- use only the math operations studied in Chapter 3.
- not use external modules, i.e., `import` command is not allowed in your solutions.

## Level 1

Level 1 problems are simple problems that required only the knowledge obtained in the content of the lecture and the reading material of the particular lecture.
For the list of chapters related to this content, please visit our syllabus in the course canvas page.

### Problem 1.1

Create a function called `sum_of_two_numbers` that receives two parameters (numbers) and returns their sum.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-1, points: 0, grade: false, solution: true -->
```python
# solution cell
# YOUR CODE HERE
raise NotImplementedError()

# uncomment the line(s) below to debug
# a = 10
# b = 20
# s = sum_of_two_numbers(a, b)
# print(f'The sum of {a} + {b} is {s}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-1-1, points: 1, grade: true, solution: false -->
```python
# test cell
from nose.tools import assert_equal, assert_raises
assert_equal(sum_of_two_numbers(2, 2), 4)
assert_equal(sum_of_two_numbers(-2, -2), -4)
assert_raises(TypeError, sum_of_two_numbers, 2) # this test makes sure that the function still requires two arguments
### THIS IS THE BEGINNING OF THE HIDDEN TESTS
assert_equal(sum_of_two_numbers(5, 2), 7)
assert_raises(TypeError, sum_of_two_numbers, 2, '1') # this makes sure that only numbers are allowed (no conversion within the function)
assert_raises(TypeError, sum_of_two_numbers, 2, None) # yet another way to test things
import random
for i in range(10):
    a = -100 + random.random() * 1000
    b = -100 + random.random() * 1000
    s = sum_of_two_numbers(a, b)
    assert_equal(s, a+b)
### THIS IS THE END OF THE HIDDEN TESTS
print('tests passed', u'\u2713')
```

### Problem 1.2

Create a function called `print_sum_of_two_numbers` that receives two parameters (numbers) and prints their sum.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-1-2, points: 0, grade: false, solution: true -->
```python
# solution cell
# YOUR CODE HERE
raise NotImplementedError()

# uncomment the line(s) below to debug
# a = 10
# b = 20
# print(f'The sum of {a} + {b} is')
# print_sum_of_two_numbers(a, b)
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-test-1-2, points: 1, grade: true, solution: false -->
```python
# test cell
from nose.tools import assert_raises
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    print_sum_of_two_numbers(10, 10)
mock_print.assert_called_once_with(20)
with patch('__main__.print') as mock_print:
    print_sum_of_two_numbers(35.5, 35)
mock_print.assert_called_once_with(70.5)
assert_raises(TypeError, print_sum_of_two_numbers, 2) # this test makes sure that the function still requires two arguments
### THIS IS THE BEGINNING OF HIDDEN TESTS
import random
for i in range(10):
    a = -100 + random.random() * 1000
    b = -100 + random.random() * 1000
    with patch('__main__.print') as mock_print:
        print_sum_of_two_numbers(a, b)
    mock_print.assert_called_once_with(a+b)
### THIS IS THE END OF HIDDEN TESTS
print('tests passed', u'\u2713')
```

---
## Level 2

Problems in level 2 should build upon the content of the lecture and the content of previous lectures.
You should connect the content related to this lecture with the content explained in previous lectures.
Therefore, the problems can be solved based solely on the reading material available for the course up to the lecture the assignment is about.

### Problem 2.1

Create a function called `print_return_sum_of_two_numbers` that receives two parameters (numbers) and, prints and returns their sum.

```
# write here your pseudocode (not graded)

# * pseudocode is necessary when asking for help during the lab
```

<!-- grade_id: cell-answer-2-1, points: 0, grade: false, solution: true -->
```python
# solution cell
# YOUR CODE HERE
raise NotImplementedError()

# uncomment the line(s) below to debug
# a = 10
# b = 20
# print(f'The sum of {a} + {b} is')
# s = print_return_sum_of_two_numbers(a, b)
# print(f'The return value is {s}')
# print('Execution finished', u'\u2713')
```

<!-- grade_id: cell-tests-2-1, points: 1, grade: true, solution: false -->
```python
# test cell
from nose.tools import assert_equal, assert_raises
from unittest.mock import patch
with patch('__main__.print') as mock_print:
    assert_equal(print_return_sum_of_two_numbers(10, 10), 20)
mock_print.assert_called_once_with(20)
with patch('__main__.print') as mock_print:
    assert_equal(print_return_sum_of_two_numbers(35.5, 35), 70.5)
mock_print.assert_called_once_with(70.5)
### THIS IS THE BEGINNING OF HIDDEN TESTS
import random
for i in range(10):
    a = -100 + random.random() * 1000
    b = -100 + random.random() * 1000
    with patch('__main__.print') as mock_print:
        assert_equal(print_return_sum_of_two_numbers(a, b), a+b)
    mock_print.assert_called_once_with(a+b)
### THIS IS THE END OF HIDDEN TESTS
print('tests passed', u'\u2713')
```

---
## Level 3: Pseudocode/HTML/CSS/JS problems

In these problems we are training the design part of the course, which will be used starting from week 6 in the course.

Recommended sources:
- HTML: https://www.w3schools.com/html/
- CSS: https://www.w3schools.com/css/
- JavaScript: https://www.w3schools.com/js/

PyCharm sometimes does not render HTML/CSS/JS as expected. We recommend testing your solutions here: https://www.w3schools.com/html/tryit.asp?filename=tryhtml_default



### Problem 3.1

Create a Python function called `get_html` that returns an HTML code that shows the Chalmers logo.

Black logo: https://www.chalmers.se/SiteCollectionImages/Logotyper/Chalmers%20logotyp/ChalmersU_black.png
White logo: https://www.chalmers.se/SiteCollectionImages/Logotyper/Chalmers%20logotyp/Chalmers_white.png

The image must have an alt text "Chalmers".

Then, make the logo be a link using the tag `a` so that when the logo is clicked, it opens the page https://student.portal.chalmers.se/en/chalmersstudies/courseinformation/Pages/SearchCourse.aspx?course_id=30582&parsergrp=3

<!-- grade_id: cell-html-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# uncomment the line(s) below to debug
# print('Your HTML:')
# print(get_html())
# print('Execution finished', u'\u2713')
```

```python
# DO NOT REMOVE THIS CODE >>> BEGINNING
# this cell assembles the HTML returned by the function that you wrote
# and renders it in the "preview" window
from IPython.core.display import display, HTML
html_header = """<!DOCTYPE html>
    <html lang="en">
        <head>
            <title>Page Title</title>
        </head>
        <body><p style="color: red;">start of your solution &gt;&gt;&gt;&gt;</p><hr/>
        """
html_header_offset = 1  # number of lines in the header
html_footer = """     <hr/><p style="color: red;">&lt;&lt;&lt;&lt; end of your solution</p></body>
</html>"""

html_to_test = html_header + get_html() + html_footer
display(HTML(html_to_test))

# DO NOT REMOVE THIS CODE <<< END
```

<!-- grade_id: cell-html-test, points: 1, grade: true, solution: false -->
```python
# test cell
import sys
import requests
from bs4 import BeautifulSoup
from utils import validate_html
has_error = False

html_to_test = html_header + get_html() + html_footer

assert html_to_test.count('\n') > 2, 'Your HTML seems to not have the sufficient lines.'

soup = validate_html(html_to_test)
tags_p = soup.find_all('p')

assert len(tags_p) == 2, f'The solution must have 9 paragraphs in your HTML. You have {len(tags_p)} instead.'

tags_a = soup.find_all('a')

assert len(tags_a) == 1, f'The HTML must have a link.'

assert tags_a[0]['href'] == 'https://student.portal.chalmers.se/en/chalmersstudies/courseinformation/Pages/SearchCourse.aspx?course_id=30582&parsergrp=3', f'The link seems to be linking to the wrong page.'

assert tags_a[0].img != None, f'The link must have an image inside of it.'

assert tags_a[0].img['alt'] == 'Chalmers', f'The alt of the image must be "Chalmers".'

print('tests passed', u'\u2713')
```